Fill in your github username and idm credentials below:

```python
github_username = ""
idm_credentials = ""
```

---

# **Applied Data Science in Medicine & Psychology - Assignment 1**
## Python Basics I
---
### **General Setup and Submission**
1. Save an editable copy of this notebook to your Google Drive ``File`` $\rightarrow$ ``Save a copy in Drive`` and remove the automatically appended "Copy of" text. **Do not change the original file name**.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** (otherwise this causes an error in your implementation) or *YOUR ANSWER HERE* with your solution.
3. If you take a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When you continue working on the assignment, do this in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your repository.
4. After finishing the assignment submit your assignment as follows:
    * Check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning.
    * Save your notebook, click ``File`` $\rightarrow$ ``Save``
    * Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
    * Select ``digikolleg-data-science-psychology/1-python-basics-I-<username>`` from the repository drop-down menu
    * Type ``1_python_basics_I.ipynb`` into the file path text box
    * Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
    * Make sure that the "Include a link to Colaboratory" box is checked
    * Press "Ok"
    * A new tab or window should open up, displaying your submitted notebook within your GitHub repository

Before the assignment deadline, you can re-submit your assignment as often as you would like. This will update your assignment file in your assignment's repository.

### **Assignment Description**
In this assignment, we cover the first part of Python basics. You will write the first lines of code, create variables, and manipulate their values using operators. Furthermore, you will learn how to display your results and play around with your first plot.

#### **Assignment Goals**
You are going to learn the following things in this assignment:
1. how to run cells and use Python and Jupyter notebooks as a calculator
2. how to display content, e.g., the values of your variables
3. how to deal with variables
4. how to use the different types of operators
5. how to use line and multi-line comments
6. that graphical visualization is done using a library and that you can change the appearance easily

#### **Assignment Point Value**
This assignment is worth a total (up to) 8 points. If your code passes all of the manual and automatic tests you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.
#### 0. Imports and Setup
Please **run** the cell below to import the libraries needed for this assignment.

```python
import matplotlib.pyplot as plt
import numpy as np
```

#### 1. Python as a Calculator - Running Cells
The easiest way to think about Python is that it works sort of like a calculator. When running a cell that contains an equation Python will calculate the value of the expression.

##### 1.1 Run the cell below and get the result of ``5 + 3``:

```python
5 + 3
```

##### 1.2 Write an equation that adds 7 and 9 and multiplies the result with 4. Run the cell to get the result:

<!-- grade_id: run_cell, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 2. Print Command
The ``print`` command is a useful function that lets you display text. Therefore, it can be used for a variety of things, e.g., displaying the results of calculations or debugging (diagnosing and correcting problems with your code implementation). One simple debugging approach is to have your program periodically print out useful messages to help you understand what's happening. The most basic way to use the ``print`` function is to call ``print(...)``, replacing the ``...`` with the text you want to be displayed, enclosed in single (``'``) or double (``"``) quotes, like this:

```
print("Hello, world!")
```

##### 2.1 Write code in the next cell to print a greeting to yourself (e.g. "Hello, Anna!"):

<!-- grade_id: print, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### 3. Variables
A variable is a named object that has a particular value. This value can be of various types such as numbers, strings, boolean (``True``/``False``) values, or objects. Variables are used to store data that can be referenced and manipulated throughout a program. They can also be reassigned to a new value as the program runs. For example, if you want to store the number of participants during a study in your program, you could create a variable called "number_participants" and assign a value to it like this:
```
number_particpiants = 24
```
If the number of participants then increases throughout the study, you can simply change the value of the variable in your program:
```
number_participants = 25
```

##### 3.1 Rewrite the calculation you did in 1.2. Create the variables ``a``, ``b``, ``c``, and ``result``. Assign 7 to ``a``, 9 to ``b``, 4 to ``c``. Rewrite the "Add 7 and 9 and multiply the result with 4" equation to "Add a and b and multiply the result with c" and assign the result to the variable ``result``:

<!-- grade_id: variables, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(result)
```

<!-- grade_id: variables_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert a
assert b
assert c
assert result # checks the correct spelling of the variable
assert result == 64, "Check your implementation, did you consider that multiplication/division is executed before addition/subtraction. You can use brackets (...) to enforce the right order."
```

#### 4. Operators
In computer programming, an operator is a symbol or keyword used to perform operations on one or more operands to produce a result. These operations can be mathematical, logical, or bitwise, depending on the type of operator used. Operators are an essential part of programming languages and allow developers to perform a wide range of tasks, from simple arithmetic operations to complex logical comparisons.
##### 4.1 Arithmetic Operators
Arithmetic Operators are used to performing basic mathematical operations.

| name | operator | example |
| :- | :-: | :- |
| plus | + | a + b |
| minus | - | a - b |
| multiplication | * | a * b |
| divison | / | a / b |
| floor divison | // | a // b |
| exponent | ** | a ** b |
| modulus | % | a % b |
| negation | - | -a |

* Create the variable ``glass_of_water`` and give it the value 3. Then add ``1`` to the variable and print the new value of the variable:

<!-- grade_id: addition, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: addition_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert glass_of_water # check the correct spelling of the variable
assert glass_of_water == 4, "Check if you assigned the increased value again to your variable glass_of_water. You can do this using the following expression: a = a + value"
```

* Create the variable ``number_apples`` and assign it to 21. Create the variable ``number_people`` and assign it to 8. Calculate the number of apples per person using the created variables, assign it to the variable ``apples_per_person``, and print the result:

<!-- grade_id: divison, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: division_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert number_apples
assert number_people
assert apples_per_person
assert apples_per_person == 2.625, "Check if you assigned the correct values to number_apples and number_people. Check if you used the right operator."
```

* Calculate ``apples_per_person`` using floor division and the variables created above. Print the result:

<!-- grade_id: floor_division, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: floor_division_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert apples_per_person
assert apples_per_person == 2, "Check if you assigned the correct values to number_apples and number_people in one of the cells above. Check if you used the right operator."
```

* Explain the difference between division and floor division. In which case could floor division be more suitable for a calculation task than division? You can use the apples_per_person example in your explanation:

*YOUR ANSWER HERE*

* Calculate ``number_apples % number_people`` using the modulus operator, assign the result to a variable called ``apples_left_over``, and print the result:

<!-- grade_id: modulus, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: modulus_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert apples_left_over
assert apples_left_over == 5, "Check if you assigned the correct values to number_apples and number_people in one of the cells above. Check if you used the right operator."
```

* Explain what a modulus operator does. In which situation could the usage of modulus be useful? You can use the apples_left_over example in your explanation:

*YOUR ANSWER HERE*

##### 4.2 Assignment Operators
Assignment operators are used to assign a value to a variable. Combinations between arithemtic and assignment operators (e.g. +=) perform the the arithemtic operation first and then assigns this manipulated value to the variable.

| name | operator | example | short for |
| :- | :-: | :- | :- |
| equals | = | a = b |  |
| plus equals | += | a += b| a = a + b |
| minus equals | -= | a -= b | a = a - b|
| multiplication equals | *= | a *= b | a = a * b
| divison equals | /= | a /= b | a = a / b |
| floor divison equals | //= | a //= b | a = a // b
| exponent equals | **= | a **= b | a = a ** b |
| modulus equals | %= | a %= b | a = a % b |

* Convert ``temp`` from degrees Fahrenheit to degrees Celsius using the following equation $temp_C = (temp_F - 32) * 5 / 9$ and only assignment operators:

<!-- grade_id: assignment_op, points: 0, grade: false, solution: true -->
```python
temp = 65.3 # degrees Fahrenheit

# YOUR CODE HERE
raise NotImplementedError()

print(temp) # print temp in degrees Celsius
```

<!-- grade_id: assignment_op_test, points: 1, grade: true, solution: false -->
```python
# run test to check your implementation
assert temp == 18.5, "Check if you used the right operators and order of operations"
```

##### 4.3 Comparison Operators
Comparison operators are used to comparing two values and return a boolean (``True``/``False``) value depending on the result of the comparison.

| name | operator | example |
| :- | :-: | :- |
| Equal To | == | a == b
| Not Equal To | != | a != b
| Greater Than | > | a > b
| Less Than | < | a < b
| Greater Than or Equal To | >= | a >= b
| Less Than or Equal To | <= | a <= b

* Given are the number of participants in study 1 ``participants_s1`` and study 2 ``participants_s2``. Create two variables:
    * ``condition_1`` checks if ``participants_s1`` is *not equal to* ``participants_s2``
    * ``condition_2`` checks if ``participants_s1`` is *greater than or equal to* ``participants_s2``

<!-- grade_id: comparison_operators, points: 0, grade: false, solution: true -->
```python
participants_s1 = 45
participants_s2 = 48

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: comparison_operators_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert condition_1 == True, "Check if you put the variables in the right order and used the correct operator for condition_1"
assert condition_2 == False,  "Check if you put the variables in the right order and used the correct operator for condition_2"
```

##### 4.4 Logical Operators
Logical operators are used to combining multiple conditions and return a boolean value based on the result of the logical operation.

<table>
<tr><th>operators</th><th>and</th><th>or</th></tr>
<tr><td>

|name|operator|example|
|:-|:-:|:-|
|and|and|a and b|
|or|or|a or b|
|not|not|a and not b|

</td><td>

|a|b|a and b| 
|:-|:-|:-|
|True|True|True|
|True|False|False|
|False|True|False|
|False|False|False|
    
</td><td>

|a|b|a or b| 
|:-|:-|:-|
|True|True|True|
|True|False|True|
|False|True|True|
|False|False|False|
</td></tr> </table>

* Create 3 variables using ``condition_1`` and ``condition_2`` from above:
    * ``result_1`` checks if ``condition_1`` and ``condition_2`` are True
    * ``result_2`` checks if ``condition_1`` or ``condition_2`` is True
    * ``result_3`` checks if ``condition_1`` and not ``condition_2`` is True

<!-- grade_id: logical_operators, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(f"{condition_1} and {condition_2} = {result_1}")
print(f"{condition_1} or {condition_2} = {result_2}")
print(f"{condition_1} and not {condition_2} = {result_3}")
```

<!-- grade_id: logical_operators_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
assert result_1 == False, "Check if you used the right operator for result_1"
assert result_2 == True, "Check if you used the right operator for result_2"
assert result_3 == True, "Check if you used the right operator for result_3"
```

* Make sure that you understood the behavior of the three logical operators!

#### 5. Comments
Comments in programming are non-executable lines of code that are added to a program to provide information or clarification about the code. They are used to help programmers understand what the code does, why it does it, or how it works. Also, it can be useful to comment parts of your implementation when looking for errors. Comments are ignored by the computer when the code is executed.

Two types of comments exist:
1. single-line comments
```
# Here you can only write in one line
```
2. multi-line (block) comments
```
"""
Here you can 
write as many lines 
as you want
"""
```

##### 5.1 Look at the code in the next cell and describe the different steps using comments:

```python
result = 8 / 5
result **= 2
print(result)
```

#### 6. Data Visualization - Plotting
Data visualization is an essential tool when working with data analysis because it helps you comprehend complex information and identify patterns that might not be immediately obvious from looking at raw data. Charts, graphs, and other graphically enhanced data can be an effective tools for presenting complex facts to others clearly and straightforwardly. Additionally, by visualizing the data, we can better identify outliers and abnormalities that may indicate errors.

In Python, the library used for generating graphics and figures is called ``matplotlib``. The library is imported in *0. Setup and Imports*. All of its functions can be accessed using ``plt.<function>``. Besides ``matplotlib``, also ``numpy`` a library for efficiently working with matrices and arrays is imported. For now, you don't need to worry about understanding the code in detail.

The following code (run the cell below) visualizes (plots) a sine signal. Besides the data itself, it is also of high importance to provide the context of the figure by including axis labels, units, title, a legend, and other relevant information.

```python
x = np.arange(start=0, stop=4*np.math.pi, step=0.1) # create a series of x values
y = np.sin(x) # sine signal

plt.plot(x, y, marker='^', label='sine') # plot signal and label for legend
plt.legend() # plot legend
plt.xlabel('x') # label the x axis
plt.ylabel('sin(x)') # label the y axis
plt.title('sine signal') # set the plot title
plt.show()
```

##### 6.1 Play around with the sine plotting code
* What part of the graph is affected if you:
1. Change start from ``0`` to ``2`` in ``np.arange(start=0, stop=4*np.math.pi, step=0.1)``.
2. Change step from ``0.1`` to ``1`` in ``np.arange(start=0, stop=4*np.math.pi, step=0.1)``.
3. Change ``np.sin(x)`` in ``y = np.sin(x)`` to ``np.cos(x)``.
4. Change ``'^'`` in ``plt.plot(x, y, marker='^', label='sine')`` to ``'s'``.

* Describe your findings for all 4 subtasks separately in the following cell:

*YOUR ANSWER HERE*